```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 3 Problem Set

## Cohort Sessions

**CS1.** You have implemented factorial using iteration in Problem Set 1. Now, implement the factorial problem using a recursion. The function should takes in an Integer input and returns and Integer output which is the factorial of the input. Recall that:

$n! = n \times (n-1) \times (n-2) \times \ldots \times 2 \times 1$

You should consider the case when $n$ is zero and one as well. 

```python
# recursion
def factorial_recursion(n):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert factorial_recursion(0) == 1
assert factorial_recursion(1) == 1
assert factorial_recursion(5) == 120
assert factorial_recursion(7) == 5040
assert factorial_recursion(11) == 39916800
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Helper Function:* Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). To do this, write another function `is_palindrome(s, left, right)` where `left` and `right` are indices from the left and from the right to check the character in `str`. Use recursion instead of iteration in this problem.


```python
def palindrome(s):
    ###
    ### YOUR CODE HERE
    ###
    pass

def is_palindrome(s, left, right):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Towers of Hanoi:* Write a function `move_disks(n, from_tower, to_tower, aux_tower)` which returns an array of String for the movement of disks that solves the [Towers of Hanoi problem](https://en.wikipedia.org/wiki/Tower_of_Hanoi). 
* The first argument `n` is an Integer input that gives information on the number of disk. 
* The second argument `from_tower` is a String which is the label of the origin tower.
* The third argument `to_tower` is a String which is the label of the destination tower.
* The last argument `aux_tower` is a String which is the label of the auxilary tower.


```python
def move_disks(n, from_tower, to_tower, aux_tower):
    result = []
    ###
    ### YOUR CODE HERE
    ###
    return result
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
result = move_disks(3, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to B.", "Move disk 2 from A to C.", "Move disk 1 from B to C.", "Move disk 3 from A to B.", "Move disk 1 from C to A.", "Move disk 2 from C to B.", "Move disk 1 from A to B."]

result = move_disks(4, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B.", "Move disk 3 from A to C.", "Move disk 1 from B to A.", "Move disk 2 from B to C.", "Move disk 1 from A to C.", "Move disk 4 from A to B.", "Move disk 1 from C to B.", "Move disk 2 from C to A.", "Move disk 1 from B to A.", "Move disk 3 from C to B.", "Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B."]

```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** *Merge Sort:* Write functions to implement the Merge Sort algorithm. The first function `mergesort(array)` should takes in an array of Integers in `array`. The function should sort the array in place. The second function `merge(array, p, q, r)` should implements the merge procedure. This function takes in an array of Integers in. `array`, the starting index for the left array `p`, the ending index for the left array `q`, and the ending index for the right array `r`. You can use a helper function for your recursion if needed.

```python
def merge(array, p, q, r):
    ###
    ### YOUR CODE HERE
    ###
    pass

def mergesort_recursive(array, p, r):
    ###
    ### YOUR CODE HERE
    ###
    pass

def mergesort(array):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
input_array = [5, 2, 4, 7, 1, 3, 2, 6]
mergesort(input_array)
assert input_array == [1, 2, 2, 3, 4, 5, 6, 7]
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```