```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 3 Problem Set

## Homeworks

**HW1.** *Fibonaci:* Write a function to find the Fibonacci number given an index. These are the example of a the first few numbers in Fibonacci series.

| Indices    | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7  | 8  | 9  |
|------------|---|---|---|---|---|---|---|----|----|----|
| The series | 0 | 1 | 1 | 2 | 3 | 5 | 8 | 13 | 21 | 34 |

We can write that the Fibonacci number at index *i* is given by:

$fibo(i) = fibo(i-1) + fibo(i-2)$

User recursion for your implementation.

```python
def fibonacci(index):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(3) == 2
assert fibonacci(7) == 13
assert fibonacci(9) == 34
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Max-Heapify:* Recall the algorithm for max-heapify in restoring the *max-heap property* of a binary heap tree. In the previous implementation, we used iteration. Implement the same algorithm using recursion. 

```python
def left_of(index):
    ###
    ### YOUR CODE HERE
    ###
    pass

def right_of(index):
    ###
    ### YOUR CODE HERE
    ###
    pass

def max_heapify(array, index, heap_size):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *String Permutation:* Write two functions to return an array of all the permutations of a string. For example, if the input is `abc`, the output should be

```
output = ["abc", "acb", "bac", "bca", "cab", "cba"]
```

The first function only has one argument which is the input string, i.e. `permutate(s)`. The second function is the recursive function with two arguments String 1 (`str1`) and and String 2 (`str2`). The first function calls the second method `permutate("", s)` at the beginning. The second function should use a loop to move a character from `str2` to `str1` and recursively invokes it with a new `str1` and `str2`. 

```python
def permutate(s):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert set(permutate("abc"))==set(["abc", "acb", "bac", "bca", "cab", "cba"])
assert set(permutate("abcd"))==set(["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"])
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *GCD:* Implement the recursive implementation of Euclid's algorithm for Greatest Common Divisor (GCD).

Hint: [Wikipedia: Greatest Common Divisor](https://en.wikipedia.org/wiki/Greatest_common_divisor#Euclid's_algorithm)

```python
def gcd(a, b):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert gcd(0,0) == 0
assert gcd(2,0) == 2
assert gcd(0,2) == 2
assert gcd(2,3) == 1
assert gcd(4, 28) == 4
assert gcd(1024, 526) == 2
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```