```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 6 Problem Set

## Cohort Sessions

**CS1.** Create a class called `Fraction` to represent a simple fraction. The class has two properties:
- `num`: which represents a numerator and of the type Integer.
- `den`: which represents a denominator and of the type Integer. Denominator should not be a zero. If a zero is assigned, you need to replace it with a `1`. 

The class should have the following method:
- `__init__(num, den)`: to initialize the numerator and the denominator. You should check if the denominator is zero. If it is you should assign `1` as the denominator instead. 
- `__str__()`:  for the object instance to be convertable to String.  You need to return a string in a format of `num/den`


```python
class Fraction:
    def __init__(self, num, den):
        self.num = num
        self.den = den
    
    @property
    def num(self):
        return self._num
    
    @num.setter
    def num(self, val):
        # if the val is not an integer, do nothing 
        if type(val) != int:
            return
        self._num = val
    
    @property
    def den(self):
        return self._den
    
    @den.setter
    def den(self, val):
        if type(val) != int or val == 0:
            self._den = 1
            return
        self._den = val
    
    def __str__(self):
        return f'{self.num}/{self.den}'
    
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
f0 = Fraction(0, 1)
assert f0.num == 0
assert f0.den == 1
assert str(f0) == "0/1"

f1 = Fraction(1, 2)
assert f1.num == 1
assert f1.den == 2
assert str(f1) == "1/2"

f1.num = 3
f1.den = 4
assert str(f1) == "3/4"
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Extend the class `Fraction` to support the following operator: `+` and `==`. To do this, you need to overload the following operator:
- `__add__(self, other)`
- `__eq__(self, other)`

You may want to write a method to simplify a fraction:
- `simplify()`: which simplify a fraction to its lowest terms. To simplify a fraction divide both the numerator and the denominator with the greatest common divisor of the the two. This method should return a new `Fraction` object.


```python
def gcd(a, b):
    if b == 0:
        return a
    else:
        return gcd(b, a % b)

class Fraction:

    def __init__(self, num, den):
        self.num = num
        self.den = den
    
    @property
    def num(self):
        return self._num
    
    @num.setter
    def num(self, val):
        # if the val is not an integer, do nothing 
        if type(val) != int:
            return
        self._num = val
    
    @property
    def den(self):
        return self._den
    
    @den.setter
    def den(self, val):
        if type(val) != int or val == 0:
            self._den = 1
            return
        self._den = val
    
    def __str__(self):
        return f'{self.num}/{self.den}'
    
    def simplify(self):
        # according to requirements
        divisor = gcd(self.num, self.den)
        return Fraction(self.num//divisor, self.den//divisor)
    
    def __add__(self, other):
        # it's math
        return Fraction( (self.num * other.den + other.num * self.den) , (self.den * other.den) ).simplify()
        
    
    def __eq__(self, other):
        f1 = self.simplify()
        f2 = other.simplify()
        return f1.num == f2.num and f1.den == f2.den

```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
f1 = Fraction(1, 2)
f2 = Fraction(2, 3)
f3 = f1 + f2 

assert str(f3) == "7/6"

f4 = Fraction(3, 5)
f5 = Fraction(1, 3)
f6 = f4 + f5 

assert str(f6) == "14/15"

f1 = Fraction(1, 2)
f2 = Fraction(2, 4)
assert f1 == f2

f3 = Fraction(2, 3)
f4 = Fraction(2, 4)
assert f3 != f4
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Inheritance:* Create a class called `MixedFraction` as a subclass of `Fraction`. A mixed fraction is a fraction that comprises of a whole number, a numerator and a denominator, e.g. `1 2/3` which is the same as `5/3`. The class has the following way of initializing its properties:
- `__init__(top, bot, whole)`: which takes in three Integers, the whole number, the numerator, and the denominator, e.g. `whole=1`, `top=2`, `bot=3`. The argument `whole` by default is `0`.  You can also specify `top` to be greater than `bot`. 

The class only has two properties:
- `num`: which is the numerator and can be greater than denominator.
- `den`: which is the denominator and must be a non-zero number.

The class should also have the following methods:
- `get_three_numbers()`: which is used to calculate the whole number, numerator and the denominator from a given numerator and denominator. The stored properties are `num` and `den` as in `Fraction` class. This function returns three Integers as a tuple, i.e. `(top, bot, whole)`.

The class should also override the `__str__()` method in this manner:
- `num/dem` if the numerator is smaller than the denominator. For example, `2/3`. 
- `whole top/bot` if the numerator is greater than the denominator. For example, `1 2/3`.

```python
class MixedFraction(Fraction):
    def __init__(self, top, bot, whole=0):
        num = top + whole * bot if whole else top
        super().__init__(num, bot)

    def get_three_numbers(self):
        # math
        whole = self.num // self.den
        top = self.num % self.den
        bot = self.den

        return (top, bot, whole)

    def __str__(self):
        if self.num > self.den:
            tup = self.get_three_numbers()
            return f"{tup[-1]} {tup[0]}/{tup[1]}"
            
        return f"{self.num}/{self.den}"
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
mf1 = MixedFraction(5, 3)
assert mf1.num == 5 and mf1.den == 3
assert mf1.get_three_numbers() == (2, 3, 1)
mf2 = MixedFraction(2, 3, 1)
assert mf2.num == 5 and mf2.den == 3

result = mf1 + mf2
assert result.num == 10 and result.den == 3

assert mf1 == mf2
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** *Inheritance:* Create a class `Deque` as a subclass of `Queue`. Use the double-stack implementation of `Queue` in this problem. `Deque` has the following methods:
- `add_front(item)`: which add an item to the front of the queue. 
- `remove_rear()`: which pops out an item from the rear of the queue. 
- `add_rear(item)`: which add an item from rear of the queue. This is the same as enqueue a normal queue.
- `remove_front()`: which pops out an item from the front of the queue. This is the same as dequeue method in a normal queue.
- `left_to_right()`: which is a helper function to move the items from the left stack to the right stack and keep the proper order. Use the Push and Pop operation of Stacks here.

See the notation for the front and rear of a Queue.

![](https://www.dropbox.com/s/xexwpbhljopmg37/queue.png?raw=1)


```python
class Stack:
    def __init__(self):
        self.ls = []
    def push(self, item):
        self.ls.append(item)
    def pop(self):
        return self.ls.pop()
    def peek(self):
        return self.ls[-1]

    @property
    def is_empty(self):
        return self.ls == []

```

```python
class Queue:
    def __init__(self):
        self.left_stack = Stack()
        self.right_stack = Stack()

    def enqueue(self, item):
        self.left_stack.push(item)

    def dequeue(self):
        if self.right_stack.is_empty:
            while not self.left_stack.is_empty:
                self.right_stack.push(self.left_stack.pop())
        return self.right_stack.pop()
```

```python
class Deque(Queue):
  
    def add_front(self, item):
        self.right_stack.push(item)
      
    def remove_front(self):
        return self.dequeue()
    
    def add_rear(self, item):
        self.enqueue(item)

    def left_to_right(self):
        if self.left_stack.is_empty:
            return
        temp = Stack()

        while not self.right_stack.is_empty:
            temp.push(self.right_stack.pop())
        while not self.left_stack.is_empty:
            self.right_stack.push(self.left_stack.pop())
        while not temp.is_empty:
            self.right_stack.push(temp.pop())
        
    def remove_rear(self):  
        self.left_to_right()
        return self.right_stack.ls.pop(0)

    def peek_front(self):
        self.left_to_right()
        return self.right_stack.ls[-1]
    
    def peek_rear(self):
        self.left_to_right()
        return self.right_stack.ls[0]
```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
q1 = Deque()
q1.add_front(1)
q1.add_front(2)
q1.add_front(3)
assert q1.remove_rear()==1
assert q1.remove_rear()==2
assert q1.remove_rear()==3

q1.add_rear(3)
q1.add_rear(2)
q1.add_rear(1)
assert q1.remove_front()==3
assert q1.remove_front()==2
assert q1.remove_front()==1

q1.add_front(1)
q1.add_rear(2)
q1.add_front(3)
q1.add_rear(4)
q1.add_front(5)
q1.add_rear(6)
assert q1.remove_rear()==6
assert q1.remove_rear()==4
assert q1.remove_rear()==2
assert q1.remove_rear()==1
assert q1.remove_rear()==3
assert q1.remove_front()==5
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS5-Prelude.** *ArrayFixedSize class (Given):* Write a class called `ArrayFixedSize` that simulate a fixed size array. This class should inherint from `collections.abc.Iterable` base class. A fixed size array is like a list which size cannot change once it is set. The size and its data type is specified during object instantiation. Use Numpy's array for its internal data storage. At the start you can use `np.empty(size)` to create an uninitalized empty array. The class should have the following methods:

- `__getitem__(index)`: which returns the element at a given `index` using the bracket operator, i.e. `array[index]`.
- `__setitem__(index, value)`: which set the item at a given `index` with a particular `value` using the bracket and assignment operators, i.e. `array[index] = value`.
- `__iter__()`: which returns the iterable object so that it can be used in a for loop and other iterable object operators.
- `__str__()`: which returns the string object representation of the object. It should displays as follows: `[el1, el2, el3, ...]`.
- `__len__()`: which returns the number of items in the array when `len()` function is called upon this object.

```python
import  collections.abc as c
import numpy as np

class ArrayFixedSize(c.Iterable):
    
    def __init__(self, size, dtype=int):
        self.__data = np.empty(size)
        self.__data = self.__data.astype(dtype)
        
    def __getitem__(self, index):
        return self.__data[index]
    
    def __setitem__(self, index, value):
        self.__data[index] = value      
        
    def __iter__(self):
        return iter(self.__data)
    
    def __len__(self):
        return len(self.__data)
        
    def __str__(self):
        out = "["
        for item in self:
            out += f"{item:}, "
        if self.__data != []:
            return out[:-2] + "]"
        else:
            return "[]"
```

**CS5.** Implement a class called `MyAbstractList` which is a subclass of Python's `collections.abc.Iterator` class. This class is meant to be an abstract class for two kinds of List data structures, i.e. `MyArrayList` and `MyLinkedList`. In this exercise, you will implement `MyAbstractList`.

This class has the following attribute and computed property property:
- `size`: which gives you the size of the list in integer.
- `is_empty`: which returns either `True` or `False` depending whether the list is empty or not.

Implement also several other methods:
- `__init__(self, list_of_items)`: which initializes the list by adding the items in list argument using the `append(item)` method.
- `__getitem__(index)`: which returns the element at a given `index` using the bracket operator, i.e. `array[index]`. This method should call the method `_get(index)` which should be implemented by the child class.
- `__setitem__(index, value)`: which set the item at a given `index` with a particular `value` using the bracket and assignment operators, i.e. `array[index] = value`. This method should call the method `_set_at(index, item)` which should be implemented in the child class.
- `__delitem__(index)`: which removes the item at a given `index` using the `del` operator, i.e. `del array[index]`. This method should call the method `_remove_at(index)` which should be implemented in the child class.
- `append(item)`: which adds an item at the end of the list. This method should call `_add_at(index, item)` implemented in the child class which adds the item at the specified index. 
- `remove(item)`: which removes the first occurence of the item in the list. This method should call `_index_of(item)` and `_remove_at(index)` implemented in the child class which removes the item at the specified index.
- `__next__()`: which returns the next element in the iterator. This method should call `_get(index)` in the child class which returns the item at the specified index. If there is no more element, it should `raise StopIteration`.

Create any other attributes that you think is necessary.





```python
import collections.abc as c

class MyAbstractList(c.Iterator):
    
    def __init__(self, list_items):
        # iterate over every element and call self.append(item)
        self.size = 0
        self._idx = 0
        for i in list_items:
            self.append(i)
    
    
    @property
    def is_empty(self):
        return self.size == 0
    
    def append(self, item):
        self._add_at(self.size, item)
        
    def remove(self, item):
        idx = self._index_of(item)
        # _index_of returns -1 if the item is not in the data, so we need to account for -1 which we are not supposed to remove
        if idx >= 0:
            self._remove_at(idx)
        
    def __getitem__(self, index):
        return self._get(index)
    
    def __setitem__(self, index, value):
        self._set_at(index, value)
        
    def __delitem__(self, index):
        self._remove_at(index)
    
    def __len__(self):
        return self.size
        
    def __iter__(self):
        self._idx = 0
        return self
        
    def __next__(self):
        if self._idx < self.size:
            n_item = self._get(self._idx)
            self._idx += 1
            return n_item
        else:
            raise StopIteration
    
    # the following methods should be implemented in the child class
    def _get(self, index):
        pass

    def _set_at(self, index, item):
        pass

    def _add_at(self, index, item):
        pass

    def _remove_at(self, index):
        pass

    def _index_of(self, item):
        pass
```

```python
# creating a class PythonList inheriting from MyAbstractList
# this is just for testing the MyAbstractList class

class PythonList(MyAbstractList):
    
    def __init__(self, init_data=[]):
        self.data = []
        super().__init__(init_data)
    
    def _add_at(self, index, item):
        self.data.insert(index, item)
        self.size += 1
        
    def _set_at(self, index, item):
        self.data[index] = item
        
    def _remove_at(self, index):
        self.data.pop(index)
        self.size -= 1
        
    def _get(self, index):
        if 0 <= index < self.size:
            return self.data[index]
        else:
            raise IndexError()
            
    def _index_of(self, item):
        try:
            idx = self.data.index(item)
            return idx
        except:
            return -1

```

<!-- grade_id: cs51, points: 1, grade: true, solution: false -->
```python
f = PythonList(list(range(10)))
# Testing init
assert f.data == list(range(10))

# Testing size property
assert f.size == 10

# Testing is_empty property
assert not f.is_empty

# Testing append method
f.append(10)
assert f.data == list(range(11))

# Testing remove method
f.remove(5)
assert f.data == [0, 1, 2, 3, 4, 6, 7, 8, 9, 10]
f._add_at(5, 5)
f.append(5)
f.remove(5)
assert f.data == [0, 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
assert not f.remove(11)

# Testing getitem method
assert [f[i] for i in range(len(f))] == [0, 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
f[0] = -1
assert [f[i] for i in range(len(f))] == [-1, 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]

# Testing delitem
del f[0]
assert f[0] == 1

# Testing __next__
assert [x for x in f] == [ 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
# Testing __iter__
assert [x for x in f] == [ 1, 2, 3, 4, 6, 7, 8, 9, 10, 5]
```

<!-- grade_id: cs52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS6.** Implement the class `MyArrayList` which is a subclass of `MyAbstractList`. `MyArrayList` should implement the `ArrayFixedSize` class instead of Python's list as in `PythonList` class above. You need to implement the following methods:
- `__init__(self, items)`: where items are used to initialize the array.
- `_add_at(self, index, item)`: which add the `item` at a particular `index`. The algorithm for this method is as follows:
    - calls `ensure_capacity()` which doubles the size of the array into a new array and copies the data.
    - shift all the items by one for every position after the `index`.
    - add the `item` at `index`.
    - add the `size` attribute by 1.
- `_set_at(self, index, value)`: this method replaces the item at `index` with a new value. 
- `_remove_at(self, index)`: which removes the item at `index`.
- `_get(self, index)`: which returns the item at `index` if the index is valid, otherwise, it raise `IndexError()`.
- `_index_of(self, item)`: which return the index of the given `item`, otherwise, it returns `-1`.
- `clear(self)`: which set the data back to zero with the size of the initial capacity.

Since we are using `ArrayFixedSize` as the internal data type, take a look at the following hints:
- We will create the array data in blocks of 16. For example, in the beginning, even if you have data less 16, we will create 16 data blocks. We will only initialize the value in that block to what we need.
- When we add data into the list, we need to ensure we have enough capacity. So when the initialized data has reach the end of 16 blocks, we need to create a new block with double the size. You then need to copy the data from the old block to the new block. 

```python
import numpy as np

class MyArrayList(MyAbstractList):
    INITIAL_CAPACITY = 16
    
    def __init__(self, items, dtype=int):
        self.data = ArrayFixedSize(MyArrayList.INITIAL_CAPACITY, dtype)
        super().__init__(items)
    
    def _add_at(self, index, item):
        self._ensure_capacity()

        for i in range(self.size-1, index-1, -1):
            # shift all the elements from element at index to last element RIGHT, meaning +1 to index to all
            # previous self.data[index] is now self.data[index+1] essentially
            self.data[i+1] = self.data[i]
        self.data[index] = item
        self.size += 1

    def _set_at(self, index, value):
        # as per ArrayFixedSize
        self.data[index] = value
        
    def _remove_at(self, index):
        if 0 <= index < self.size:
            res = self.data[index]
            for idx in range(index, self.size):
                # shift all the elements from element at index+1 to last element LEFT, meaning -1 index to all
                self.data[idx] = self.data[idx+1]
            self.size -= 1
            return res
         
        else:
            raise IndexError()
        
    def _get(self, index):
        if 0 <= index < self.size:
            return self.data[index]
        else:
            raise IndexError()
            
    def _index_of(self, item):
        for i in range(self.size):
            if self.data[i] == item:
                return i
        return -1
        
    def _ensure_capacity(self):
        if self.size >= len(self.data):
            new_data = ArrayFixedSize(self.size * 2 + 1)
            self._copy(self.data, 0, new_data, 0)
            self.data = new_data
            
    def _copy(self, source, idx_s, dest, idx_d):
        for idx in range(idx_s,len(source)):
            offset = idx - idx_s
            dest[idx_d + offset] = source[idx]
            
    def _clear(self):
        self.data = ArrayFixedSize(MyArrayList.INITIAL_CAPACITY)
        self.size = 0
        
    def __str__(self):
        out = "["
        for idx in range(self.size):
            out += f"{self._get(idx):}, "
        return out[:-2] + "]"

    # MISUNDERSTOOD WHAT COPY WANTED
    # new_data = ArrayFixedSize(len(self.data))
    # new_data[index] = item
    # for i in range(index):
    #     new_data[i] = self.data[i]
    # for i in range(index+1, self.size):
    #     new_data[i] = self.data[i-1]
    # self.data = new_data
    # self.size += 1
```

<!-- grade_id: cs61, points: 1, grade: true, solution: false -->
```python
a = MyArrayList([1,2,3])
assert [x for x in a] == [1,2,3]
assert a.size == 3
assert not a.is_empty
a.append(4)
assert a.size == 4
assert [x for x in a] == [1,2,3,4]
assert [a[i] for i in range(len(a))] == [1,2,3,4]
a[0] = -1
assert [x for x in a] == [-1, 2,3,4]
del a[0]
assert [x for x in a] == [2,3,4]
```

<!-- grade_id: cs62, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```