# Week 8 Problem Set

## Homeworks

```python
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
import numpy as np
```

**HW1.** *Breast Cancer Data:* Read Breast Cancer data using Pandas library:
- **Task 1:** Read the file: `breast_cancer_data.csv`:
    - [Download the CSV file](https://www.dropbox.com/s/1u8y8kzdltse4yd/data.csv?raw=1)
    - [Read Description of Data](https://www.kaggle.com/uciml/breast-cancer-wisconsin-data)
    - use comma `","` as field separator
    - use `"utf-8"` as encoding field

```python
# Task 1
# Read file, replace the None
df = pd.read_csv("breast_cancer_data.csv", sep=",", encoding="utf-8")

###
### YOUR CODE HERE
###
display(df)

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
assert isinstance(df, pd.DataFrame)
assert df.shape == (569, 32)
assert df.columns[0] == 'id' and df.columns[-1] == 'fractal_dimension_worst'
```

- **Task 2:** Find the number of rows and columns.


```python
# Task 2
# get the shape
shape = df.shape

# get rows and columns from shape
row = shape[0]
col = shape[1]

###
### YOUR CODE HERE
###

print(row,col)


```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
assert shape == (569, 32)
assert row == 569
assert col == 32
```

- **Task 3:** Find the name of all the columns.


```python
# Task 3
# display the name of all the columns
names = df.columns

###
### YOUR CODE HERE
###

print(names)
```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
assert isinstance(names, pd.Index)
assert np.all(names == pd.Index(['id', 'diagnosis', 'radius_mean', 'texture_mean', 'perimeter_mean',
       'area_mean', 'smoothness_mean', 'compactness_mean', 'concavity_mean',
       'concave points_mean', 'symmetry_mean', 'fractal_dimension_mean',
       'radius_se', 'texture_se', 'perimeter_se', 'area_se', 'smoothness_se',
       'compactness_se', 'concavity_se', 'concave points_se', 'symmetry_se',
       'fractal_dimension_se', 'radius_worst', 'texture_worst',
       'perimeter_worst', 'area_worst', 'smoothness_worst',
       'compactness_worst', 'concavity_worst', 'concave points_worst',
       'symmetry_worst', 'fractal_dimension_worst']))
```

- **Task 4:** Create a subset data set containing only the following features:
    - radius (mean of distances from center to points on the perimeter)
    - texture (standard deviation of gray-scale values)
    - perimeter
    - area
    - smoothness (local variation in radius lengths)
    - concavity (severity of concave portions of the contour)
    
**Make sure the data type is `pd.DataFrame`.**


```python
# Task 4
# set the name of the columns for the subset of data
columns = ['radius_mean', 'texture_mean', 'perimeter_mean', 'area_mean', 'smoothness_mean', 'concavity_mean']

# extract the columns, replace the None
df_features = df.loc[:, columns]


```

<!-- grade_id: hw14, points: 1, grade: true, solution: false -->
```python
assert isinstance(df_features, pd.DataFrame)
assert df_features.columns[0] == 'radius_mean'
assert df_features.columns[-1] == 'concavity_mean'
assert df_features.shape == (569, 6)
```

- **Task 5:** Create a subset data set containing only the target from the column "diagnosis". 

**Make sure the data type is `pd.DataFrame`.**

```python
#5.
# extract target
df_target = pd.DataFrame(df.loc[:, 'diagnosis'])

###
### YOUR CODE HERE
###

df_target
```

<!-- grade_id: hw15, points: 1, grade: true, solution: false -->
```python
assert isinstance(df_target, pd.DataFrame)
assert df_target.shape == (569, 1)
assert df_target.columns[0] == 'diagnosis'
```

- **Task 6:** Create a new Data Frame from the column "diagnosis", called "diagnosis_int" which is the integer representation of the column diagnosis. Copy the column into the original data frame so that it has a copy.
    - if the value in "diagnosis" column is "M", the value in the column "diagnosis_int" should be set to 1
    - otherwise, it should be set to 0
    
*Hint: use `.apply()` method.*


```python
# Task 6
# creating diagnosis_int
df_target["diagnosis_int"] = df_target['diagnosis'].apply(lambda x: 1 if x == "M" else 0)

# copy the new column into the original data frame
df["diagnosis_int"] = np.copy(df_target["diagnosis_int"])

###
### YOUR CODE HERE
###

display(df_target)
display(df["diagnosis_int"])
```

<!-- grade_id: hw16, points: 1, grade: true, solution: false -->
```python
assert isinstance(df_target, pd.DataFrame) and isinstance(df, pd.DataFrame)
assert df_target.shape == (569, 2)
assert np.all(df_target.columns == ["diagnosis", "diagnosis_int"])
assert "diagnosis_int" in df.columns
```

- **Task 7:** Use scatter plot to see the relationship between "radius_mean" and "diagnosis_int". 
    - Use the "diagnosis" column as the `hue` for the scatter plot.
    - Label the x axis as "Mean of Cell Radius"
    - Label the y axis as "1: Malignant, 0: Benign"

```python
# Task 7
# set the default theme to use Seaborn
sns.set()

# display using scatter plot
myplot = sns.scatterplot(data=df, x="radius_mean", y="diagnosis_int", hue="diagnosis")

# set the x label
myplot.set_xlabel("Mean of Cell Radius")
myplot.set_ylabel("1: Malignant, 0: Benign")
# set the y label

###
### YOUR CODE HERE
###

```

- **Task 8:** Use scatter plot to see the relationship between "concavity_mean" and "diagnosis_int". 
    - Use the "diagnosis" column as the `hue` for the scatter plot.
    - Label the x axis as "Mean of Cell Concavity"
    - Label the y axis as "1: Malignant, 0: Benign"

```python
# Task 8
# set the default theme to use Seaborn
sns.set()

# display using scatter plot
myplot = sns.scatterplot(data=df, x="concavity_mean", y="diagnosis_int", hue="diagnosis")

# set the x label
myplot.set_xlabel("Mean of Cell Concavity")
myplot.set_ylabel("1: Malignant, 0: Benign")

# set the y label

###
### YOUR CODE HERE
###

```

**HW2.** *Count Plot:* Create a function to count how many records are diagnosed as Malignant and Benign. The function should return a tuple: `(n_malignant, n_benign)`, where `n_malignant` is the number of records diagnosed as Malignant cell and `n_benign` is the number of records diagnosed as Benign.

Use Count plot to verify the answer.

Reference:
- [Count Plot](https://seaborn.pydata.org/generated/seaborn.countplot.html)


```python
def count_cell(target):
    n_malignant = 0
    n_benign = 0
    for cell in target:
        if cell == "M":
            n_malignant += 1
        else:
            n_benign += 1
        
    return (n_malignant, n_benign)

```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
result = count_cell(df_target["diagnosis"])
assert result == (212, 357)
```

```python
# write the code to plot the count of the two classes
myplot = sns.countplot(data=df_target, x='diagnosis')

###
### YOUR CODE HERE
###

```

**HW3.** *Normalization:* Create a function that takes in Data Frame as the input and returns the normalized Data Frame as the output. **Each column** is normalized separately using **min-max normalization**. The function should return a new data frame instead of modifying the input data frame.

$$normalized = \frac{data - min}{max - min}$$

Use the following functions from Pandas or Numpy:
- `df.copy()`: This is to create a new copy of the data frame.
- `df.min(axis=0)`: This is to get the minimum along the index axis.
- `df.max(axis=0)`: This is to get the maximum along the index axis.

Note: Your function should be able to handle a numpy array as well as Panda's data frame. 

```python
def normalize_minmax(dfin):
    dfout = dfin.copy()
    dfout_min = dfout.min(axis=0)
    normalized_data = (dfout - dfout_min)/(dfout.max(axis=0)-dfout_min)
    return normalized_data
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
data_norm = normalize_minmax(df_features)
stats = data_norm.describe()
display(stats)
assert stats.loc["max", "radius_mean"] == 1.0 and \
       stats.loc["min", "radius_mean"] == 0 and \
       np.isclose(stats.loc["mean", "radius_mean"], 0.338222)
assert stats.loc["max", "texture_mean"] == 1.0 and \
       stats.loc["min", "texture_mean"] == 0 and\
       np.isclose(stats.loc["mean", "texture_mean"], 0.323965)

```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
data_norm = normalize_minmax(df_features.to_numpy())
assert data_norm[:,0].max() == 1.0 and \
       data_norm[:,0].min() == 0 and \
       np.isclose(data_norm[:,0].mean(), 0.338222)
assert data_norm[:,1].max() == 1.0 and \
       data_norm[:,1].min() == 0 and\
       np.isclose(data_norm[:,1].mean(), 0.323965)

```

**HW4.** *Splitting the Data:* Use the function to split the breast cancer data set into a training data set and a testing data set. Use `random_state=100` and `test_size=0.3` and the normalized features data set from the previous exercise.

```python
# copy your answer from the cohort problem

def split_data(df_feature, df_target, random_state=None, test_size=0.5):
    indexes = df_feature.index
    k = int(len(indexes) * test_size)
    if random_state:
        np.random.seed(random_state)
    test_indexes = set(np.random.choice(indexes, k, replace=False))
    train_indexes = set(indexes) - test_indexes
    df_feature_train = df_feature.loc[list(train_indexes), :]
    df_feature_test = df_feature.loc[list(test_indexes), :]
    df_target_train = df_target.loc[list(train_indexes), :]
    df_target_test = df_target.loc[list(test_indexes), :]
    return df_feature_train, df_feature_test, df_target_train, df_target_test
 
```

```python
# call normalize_minmax() to normalize df_features
data_norm = normalize_minmax(df_features)

###
### YOUR CODE HERE
###

# call split_data() with seed of 100 and test size of 30%
datasets_tupple = split_data(data_norm, df_target, random_state=100, test_size=0.3)
###
### YOUR CODE HERE
###


```

```python
df_features_train = datasets_tupple[0]
df_features_test = datasets_tupple[1]
df_target_train = datasets_tupple[2]
df_target_test = datasets_tupple[3]

display(df_features_train.describe())
display(df_features_test.describe())
display(df_target_train.describe())
display(df_target_test.describe())
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert isinstance(df_features_train, pd.DataFrame)
assert isinstance(df_features_test, pd.DataFrame)
assert isinstance(df_target_train, pd.DataFrame)
assert isinstance(df_target_test, pd.DataFrame)

assert df_features_train.shape == (399, 6)
assert df_features_test.shape == (170, 6)
assert df_target_train.shape == (399, 2)
assert df_target_test.shape == (170, 2)

assert np.isclose(df_features_train.mean().mean(), 0.29844) 
assert np.isclose(df_features_test.mean().mean(), 0.311966) 
assert np.isclose(df_target_train.mean().mean(), 0.358396) 
assert np.isclose(df_target_test.mean().mean(), 0.40588) 
```

**HW5.** *Pair Plot:* Use pair plot to find out the relationship between different columns in `df_features`. Ensure that similar relationship exists in both the training and the test datasets.

```python
# write your code below to plot for df_features
myplot = sns.pairplot(data=df_features)

###
### YOUR CODE HERE
###

```

```python
# write your code below to plot for df_features_train
myplot = sns.pairplot(data=df_features_train)

###
### YOUR CODE HERE
###

```

```python
# write your code below to plot for df_features_test
myplot = sns.pairplot(data=df_features_test)

###
### YOUR CODE HERE
###

```

```python

```