```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 3 Problem Set

## Cohort Sessions

**CS1.** You have implemented factorial using iteration in Problem Set 1. Now, implement the factorial problem using a recursion. The function should takes in an Integer input and returns and Integer output which is the factorial of the input. Recall that:

$n! = n \times (n-1) \times (n-2) \times \ldots \times 2 \times 1$

You should consider the case when $n$ is zero and one as well. 

```python
# recursion
def factorial_recursion(n):
    if n == 1 or n == 0:
        return 1
    return factorial_recursion(n - 1) * n

```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert factorial_recursion(0) == 1
assert factorial_recursion(1) == 1
assert factorial_recursion(5) == 120
assert factorial_recursion(7) == 5040
assert factorial_recursion(11) == 39916800
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Helper Function:* Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). To do this, write another function `is_palindrome(s, left, right)` where `left` and `right` are indices from the left and from the right to check the character in `str`. Use recursion instead of iteration in this problem.


```python
# not sure why they asked for a helper function version but it's ok we make do

def palindrome(s):
    if len(s) < 2:
        return True
    # indices left and right passed into is_palindrome
    return is_palindrome(s, s[0], s[-1])

def is_palindrome(s, left, right):
    # if left equal to right, return string without first and last char, else false
    if left == right:
        return palindrome(s[1:-1])
    return False

```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Towers of Hanoi:* Write a function `move_disks(n, from_tower, to_tower, aux_tower)` which returns an array of String for the movement of disks that solves the [Towers of Hanoi problem](https://en.wikipedia.org/wiki/Tower_of_Hanoi). 
* The first argument `n` is an Integer input that gives information on the number of disk. 
* The second argument `from_tower` is a String which is the label of the origin tower.
* The third argument `to_tower` is a String which is the label of the destination tower.
* The last argument `aux_tower` is a String which is the label of the auxilary tower.


```python
def move_disks(n, from_tower, to_tower, aux_tower):
    result = []
    # the nth to 1st iteration is really just a bunch of 1 disk movements, so base case is 1
    if n == 1:
        result.append(gen_string(n, from_tower, to_tower))
        return result
    # unmove the disk 
    result += (move_disks(n-1, from_tower, aux_tower, to_tower))
    # move nth disk
    result.append(gen_string(n, from_tower, to_tower))
    # move the rest of disks
    result += (move_disks(n-1, aux_tower, to_tower, from_tower))
    return result
def gen_string(n, t1, t2):
    # just to make my life easier
    return "Move disk {} from {} to {}.".format(n, t1, t2)
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
result = move_disks(3, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to B.", "Move disk 2 from A to C.", "Move disk 1 from B to C.", "Move disk 3 from A to B.", "Move disk 1 from C to A.", "Move disk 2 from C to B.", "Move disk 1 from A to B."]

result = move_disks(4, "A", "B", "C")
print(result)
assert result == ["Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B.", "Move disk 3 from A to C.", "Move disk 1 from B to A.", "Move disk 2 from B to C.", "Move disk 1 from A to C.", "Move disk 4 from A to B.", "Move disk 1 from C to B.", "Move disk 2 from C to A.", "Move disk 1 from B to A.", "Move disk 3 from C to B.", "Move disk 1 from A to C.", "Move disk 2 from A to B.", "Move disk 1 from C to B."]

```

Output:
```
['Move disk 1 from A to B.', 'Move disk 2 from A to C.', 'Move disk 1 from B to C.', 'Move disk 3 from A to B.', 'Move disk 1 from C to A.', 'Move disk 2 from C to B.', 'Move disk 1 from A to B.']
['Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.', 'Move disk 3 from A to C.', 'Move disk 1 from B to A.', 'Move disk 2 from B to C.', 'Move disk 1 from A to C.', 'Move disk 4 from A to B.', 'Move disk 1 from C to B.', 'Move disk 2 from C to A.', 'Move disk 1 from B to A.', 'Move disk 3 from C to B.', 'Move disk 1 from A to C.', 'Move disk 2 from A to B.', 'Move disk 1 from C to B.']
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** *Merge Sort:* Write functions to implement the Merge Sort algorithm. The first function `mergesort(array)` should takes in an array of Integers in `array`. The function should sort the array in place. The second function `merge(array, p, q, r)` should implements the merge procedure. This function takes in an array of Integers in. `array`, the starting index for the left array `p`, the ending index for the left array `q`, and the ending index for the right array `r`. You can use a helper function for your recursion if needed.

```python
# divide and conquer using index values, saves a lot of time on copying over arrays and splitting arrays
def mergesort(array):
    mergesort_recursive(array, 0, len(array))

# renamed variables for clarity
def mergesort_recursive(array, start, end):
    # if "array" is not length 1, pseudo divide by indexes
    if end - start >= 2:
        midpoint = (start + end) // 2
        # sort out the left and right halves
        mergesort_recursive(array, start, midpoint)
        mergesort_recursive(array, midpoint, end)

        # call merge on both sorted halves
        merge(array, start, midpoint, end)
    

def merge(array, start, left, right):
    # create copies of left and right
    left_array = array[start:left]
    right_array = array[left:right]

    # index pointers for left and right array
    left_index = right_index = 0
    
    # left and right size, advantage of index manipulation is that you don't need to call len(left) and len(right), which is O(n) time wasters everytime it's called
    left_size = left - start
    right_size = right - left

    # while still able to compare pointer values with left and right arrays: 
    while left_index < left_size and right_index < right_size:

        # looking for minimum value as we are sorting in ascending order
        if left_array[left_index] > right_array[right_index]:
            min_value = right_array[right_index]
            right_index += 1
        else:
            min_value = left_array[left_index]
            left_index += 1

        # replace value at main_array index with min_value
        array[start] = min_value
        start += 1
    
    # add the remainding arrays if any
    while left_index < left_size:
        array[start] = left_array[left_index]
        left_index += 1
        start += 1

    while right_index < right_size:
        array[start] = right_array[right_index]
        right_index += 1
        start += 1

```

<!-- grade_id: cs41, points: 1, grade: true, solution: false -->
```python
input_array = [5, 2, 4, 7, 1, 3, 2, 6]
mergesort(input_array)
print(input_array)
assert input_array == [1, 2, 2, 3, 4, 5, 6, 7]
```

Output:
```
[1, 2, 2, 3, 4, 5, 6, 7]
```

<!-- grade_id: cs42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```