```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 3 Problem Set

## Homeworks

**HW1.** *Fibonaci:* Write a function to find the Fibonacci number given an index. These are the example of a the first few numbers in Fibonacci series.

| Indices    | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7  | 8  | 9  |
|------------|---|---|---|---|---|---|---|----|----|----|
| The series | 0 | 1 | 1 | 2 | 3 | 5 | 8 | 13 | 21 | 34 |

We can write that the Fibonacci number at index *i* is given by:

$fibo(i) = fibo(i-1) + fibo(i-2)$

User recursion for your implementation.

```python
# best is functools.cache
def cache(func):
    memory = {}
    # now any argument passed through fibonacci is first passed through wrap, wrap head around this
    def wrap(*args):
        if args not in memory:
            memory[args] = func(*args)
        return memory[args]
    return wrap

@cache
def fibonacci(index):
    # base case is 0 and 1
    if index == 0:
        return 0
    elif index == 1:
        return 1
    # there are actually many repeaters
    return fibonacci(index - 1) + fibonacci(index - 2)


```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
assert fibonacci(0) == 0
assert fibonacci(1) == 1
assert fibonacci(3) == 2
assert fibonacci(7) == 13
assert fibonacci(9) == 34
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** *Max-Heapify:* Recall the algorithm for max-heapify in restoring the *max-heap property* of a binary heap tree. In the previous implementation, we used iteration. Implement the same algorithm using recursion. 

```python
def left_of(index):
    return index * 2 + 1

def right_of(index):
    return (index + 1) * 2


def max_heapify(array, index, heap_size):
    left = left_of(index)
    # if no children, quit recursion
    if left >= heap_size:
        return
    right = right_of(index)
    # if only left child exists
    if right >= heap_size and left < heap_size:
        max_child = left
    else:
        # ternary operator to look nice, picks out index of max child
        max_child = left if array[left] > array[right] else right
    # max heap: if child > parent, swap
    if array[max_child] > array[index]:
        array[max_child], array[index] = array[index], array[max_child]
        max_heapify(array, max_child, heap_size)


```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
result = [4, 1, 10, 14, 16, 9, 3, 2, 8, 7]
max_heapify(result, 1, len(result))
print(result)
assert result == [4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[4, 16, 10, 14, 7, 9, 3, 2, 8, 1]
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** *String Permutation:* Write two functions to return an array of all the permutations of a string. For example, if the input is `abc`, the output should be

```
output = ["abc", "acb", "bac", "bca", "cab", "cba"]
```

The first function only has one argument which is the input string, i.e. `permutate(s)`. The second function is the recursive function with two arguments String 1 (`str1`) and and String 2 (`str2`). The first function calls the second method `permutate("", s)` at the beginning. The second function should use a loop to move a character from `str2` to `str1` and recursively invokes it with a new `str1` and `str2`. 

```python
# strings are immutable, try do some cocatenates and lists
'''
def permutate(s):
    ls = []
    # assumption that a string is only 1 word max
    s = [char for char in s]
    
    # size 
    size = len(s)
    permute(s, size, ls)
    return ls

def permute(s, size, ls):
    # heap's algorithm - interesting time complexity https://stackoverflow.com/questions/42877789/heaps-algorithm-time-complexity
    if size == 1:
        ls.append("".join(s))
    else:
        for i in range(size):
            permute(s, size - 1, ls)
            if size % 2 == 0:
                s[i], s[size - 1] = s[size - 1], s[i]
            else:
                s[0], s[size - 1] = s[size - 1], s[0]
'''
# my inferior algorithm
def permutate(s):
    ls = []
    permute("", s, ls)
    return ls

def permute(str1, str2, ls):
    # stop recursion if str1 has formed length of str2
    if len(str1) == len(str2):
        ls.append(str1)
    else:
        # very very inefficient
        for char in str2:
            if not char in str1:
                permute(str1 + char, str2, ls)


```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
assert set(permutate("abc"))==set(["abc", "acb", "bac", "bca", "cab", "cba"])
assert set(permutate("abcd"))==set(["abcd", "abdc", "acbd", "acdb", "adbc", "adcb", "bacd", "badc", "bcad", "bcda", "bdac", "bdca", "cabd", "cadb", "cbad", "cbda", "cdab", "cdba", "dabc", "dacb", "dbac", "dbca", "dcab", "dcba"])
```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** *GCD:* Implement the recursive implementation of Euclid's algorithm for Greatest Common Divisor (GCD).

Hint: [Wikipedia: Greatest Common Divisor](https://en.wikipedia.org/wiki/Greatest_common_divisor#Euclid's_algorithm)

```python
def gcd(a, b):
    # if a=b, return according to algo
    if a == b:
        return a
    # if either a or b is 0, return largest value
    elif not a or not b:
        return max(a, b)
    # algo, ternary
    return gcd(a - b, b) if a > b else gcd(a, b - a) 
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
assert gcd(0,0) == 0
assert gcd(2,0) == 2
assert gcd(0,2) == 2
assert gcd(2,3) == 1
assert gcd(4, 28) == 4
assert gcd(1024, 526) == 2
```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```