```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 4 Problem Set

## Homeworks

**HW1.** Implement `Queue` abstract data structure using a Class. You can use a `list` as its internal data structure. The class should have the following interface:
- `__init__()` to initialize an empty List for the queue to store the items.
- `enqueue(item)` which inserts an Integer into the queue.
- `dequeue()` which returns and removes the element at the head of the queue. The return value is an optional as it may return `None` if there are no more elements in the queue.
- `peek()` which returns the element at the head of the queue.

The class Queue has two computed properties:
- `is_empty` which returns either `True` or `False` depending on whether the queue is empty or not.
- `size` which returns the number of items in the queue.

```python
class Queue:
    def __init__(self):
        self.__items = []
    
    def enqueue(self, item):
        self.__items.append(item)
    
    def dequeue(self):
        return self.__items.pop(0)
    
    def peek(self):
        return self.__items[0]
    
    @property
    def is_empty(self):
        return self.__items == []

    @property
    def size(self):
        return len(self.__items)
```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
q1 = Queue()
q1.enqueue(2)
assert not q1.is_empty
assert q1.size == 1
ans = q1.dequeue()
assert ans == 2
assert q1.is_empty
q1.enqueue(1)
q1.enqueue(2)
q1.enqueue(3)
assert q1.size == 3
assert q1.peek() == 1
assert q1.dequeue() == 1
assert q1.dequeue() == 2
assert q1.dequeue() == 3
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

<!-- grade_id: hw13, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** We are going to create a class that contains both `RobotTurtle` and `Coordinate` class. The name of the class is `TurtleWorld` which is used to simulate when `RobotTurtle` is moving around some two dimensional space. The class has the following methods:

- `add_turtle(name)` which is to add a new `RobotTurtle` into the world with the specified name.
- `remove_turtle(name)` which is to remove the object `RobotTurtle` with the specified name from the world. 
- `list_turtles()` which is to list all the turtles in the world using their names in an ascending order.

We give you here the class definition for the `Coordinate` and the `RobotTurtle` from the Notes.

```python
import math

class Coordinate:
    
    def __init__(self, x=0, y=0):
        self.x = x
        self.y = y
        
    @property
    def distance(self):
        return math.sqrt(self.x * self.x + self.y * self.y)
    
    def __str__(self):
        return f"({self.x}, {self.y})"
```

```python
# Class definition
class RobotTurtle:
    # Attributes:
    def __init__(self, name, speed=1):
        self.name = name
        self.speed = speed
        self._pos = Coordinate(0, 0)
        
    # property getter
    @property
    def name(self):
        return self._name
    
    # property setter
    @name.setter
    def name(self, value):
        if isinstance(value, str) and value != "":
            self._name = value
            
    # property getter
    @property
    def speed(self):
        return self._speed
    
    # property setter
    @speed.setter
    def speed(self, value):
        if isinstance(value, int) and value > 0:
            self._speed = value

    # property getter
    @property
    def pos(self):
        return self._pos
    
    # Methods:
    def move(self, direction):
        update = {'up' : Coordinate(self.pos.x, self.pos.y + self.speed),
                  'down' : Coordinate(self.pos.x, self.pos.y - self.speed),
                  'left' : Coordinate(self.pos.x - self.speed, self.pos.y),
                  'right' : Coordinate(self.pos.x + self.speed, self.pos.y)}
        self._pos = update[direction]

        
    def tell_name(self):
        print(f"My name is {self.name}")

```

Now fill in the class definition for `TurtleWorld`. You may want to look into the test cases in the following cell to make sure you define the class properly.

```python
'''
- `add_turtle(name)` which is to add a new `RobotTurtle` into the world with the specified name.
- `remove_turtle(name)` which is to remove the object `RobotTurtle` with the specified name from the world. 
- `list_turtles()` which is to list all the turtles in the world using their names in an ascending order.
'''

class TurtleWorld:
    
    def __init__(self):
        self.turtles = {}
        
    def add_turtle(self, name, speed):
        # append to self.turtles a RobotTurtle obj
        self.turtles[name] = RobotTurtle(name, speed)
        
    def remove_turtle(self, name):
        del self.turtles[name]
        
    def list_turtles(self):
        # list comprehension one-liner, sorted() to return a sorted list in ascending order
        return sorted([name for name in self.turtles.keys()])
```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
world = TurtleWorld()
world.add_turtle('t1', 1)
assert world.list_turtles() == ['t1']

world.add_turtle('t2', 2)
assert world.list_turtles() == ['t1', 't2']

world.add_turtle('abc', 3)
assert world.list_turtles() == ['abc', 't1', 't2']

world.remove_turtle('t2')
assert world.list_turtles() == ['abc', 't1']

world.remove_turtle('abc')
assert world.list_turtles() == ['t1']
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW3.** Modify the class `TurtleWorld` to add the following method:
- `move_turtle(name, movement)` which is to move the turtle with the specified name with a given input `movement`. The argument `movement` is a string containing letters: `l` for left, `r` for right, `u` for up, and `d` for down. The movement should be based on the speed. This means that if the turtle has speed of 2 and the `movement` argument is `uulrdd`, the turtle should move up four units, move left two units, move right two units and move down four units.

```python
class TurtleWorld:
    valid_movements = set('udlr')
    movement_map = {'u': 'up', 'd': 'down', 'l': 'left', 'r': 'right'}
    
    def __init__(self):
        self.turtles = {}
        
    def move_turtle(self, name, movement):
        # for each move that is in movement_map, do RobotTurtle.move(direction)
        for move in movement:
            if move in self.valid_movements:
                self.turtles[name].move(self.movement_map[move])

    def add_turtle(self, name, speed):
        # append to self.turtles a RobotTurtle obj
        self.turtles[name] = RobotTurtle(name, speed)
        
    def remove_turtle(self, name):
        del self.turtles[name]
        
    def list_turtles(self):
        # list comprehension one-liner, sorted() to return a sorted list in ascending order
        return sorted([name for name in self.turtles.keys()])
```

<!-- grade_id: hw31, points: 1, grade: true, solution: false -->
```python
world = TurtleWorld()
world.add_turtle('abc', 1)
world.move_turtle('abc', 'uu')
assert str(world.turtles['abc'].pos) == '(0, 2)'

world.move_turtle('abc', 'rrr')
assert str(world.turtles['abc'].pos) == '(3, 2)'

world.move_turtle('abc', 'd')
assert str(world.turtles['abc'].pos) == '(3, 1)'

world.move_turtle('abc', 'llll')
assert str(world.turtles['abc'].pos) == '(-1, 1)'

world.add_turtle('t1', 2)
world.move_turtle('t1', 'uulrdd')
assert str(world.turtles['t1'].pos) == '(0, 0)'

world.move_turtle('t1', 'ururur')
assert str(world.turtles['t1'].pos) == '(6, 6)'



```

<!-- grade_id: hw32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW4.** Modify the class `TurtleWorld` to include the following method:
- `add_movement(turtle, movement)` which adds turtle movement to a queue to be run later. The argument `turtle` is a string containing the turtle's name. The argument `movement` is another string for the movement. For example, value for `turtle` can be something like `'t1'` while the value for the `movement` can be something like `'uullrrdd'`.
- `run()` which executes all the movements in the queue.

```python
# newRobotTurtle extends RobotTurtle, to add queue obj to newRobotTurtle
class newRobotTurtle(RobotTurtle):
    def __init__(self, *args):
        RobotTurtle.__init__(self, *args)
        self.movement_queue = Queue()

class TurtleWorld:
    valid_movements = set('udlr')
    movement_map = {'u': 'up', 'd': 'down', 'l': 'left', 'r': 'right'}
    
    def __init__(self):
        self.turtles = {}
        
    def add_movement(self, turtle, movement):
        # enqueue movement to individual turtle
        self.turtles[turtle].movement_queue.enqueue(movement)
    
    def run(self):
        # for turtle obj in self.turtles, if movement queue is not empty, while movement queue is not empty, move_turtle()
        for turtle in self.turtles.values():
            if not turtle.movement_queue.is_empty:
                while not turtle.movement_queue.is_empty:
                    self.move_turtle(turtle.name, turtle.movement_queue.dequeue())
        
    def move_turtle(self, name, movement):
        # for each move that is in movement_map, do RobotTurtle.move(direction)
        for move in movement:
            if move in self.valid_movements:
                self.turtles[name].move(self.movement_map[move])

    def add_turtle(self, name, speed):
        # append to self.turtles a RobotTurtle obj and Queue obj 
        self.turtles[name] = newRobotTurtle(name, speed)
        
    def remove_turtle(self, name):
        del self.turtles[name]
        
    def list_turtles(self):
        # list comprehension one-liner, sorted() to return a sorted list in ascending order
        return sorted([name for name in self.turtles.keys()])
        
```

<!-- grade_id: hw41, points: 1, grade: true, solution: false -->
```python
world = TurtleWorld()
world.add_turtle('t1', 1)
world.add_turtle('t2', 2)
world.add_movement('t1', 'ur')
world.add_movement('t2', 'ur')
assert str(world.turtles['t1'].pos) == '(0, 0)'
assert str(world.turtles['t2'].pos) == '(0, 0)'

world.run()
assert str(world.turtles['t1'].pos) == '(1, 1)'
assert str(world.turtles['t2'].pos) == '(2, 2)'

world.run()
assert str(world.turtles['t1'].pos) == '(1, 1)'
assert str(world.turtles['t2'].pos) == '(2, 2)'


```

<!-- grade_id: hw42, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW5.** Implement a radix sorting machine. A radix sort for base 10 integers is a *mechanical* sorting technique that utilizes a collection of bins:
- one main bin 
- 10 digit-bins

Each bin acts like a *queue* and maintains its values in the order that they arrive. The algorithm works as follows:
- it begins by placing each number in the main bin. 
- Then it considers each value digit by digit. The first value is removed from the main bin and placed in a digit-bin corresponding to the digit being considered. For example, if the ones digit is being considered, 534 will be placed into digit-bin 4 and 667 will placed into digit-bin 7. 
- Once all the values are placed into their corresponding digit-bins, the values are collected from bin 0 to bin 9 and placed back in the main bin (in that order). 
- The process continues with the tens digit, the hundreds, and so on. 
- After the last digit is processed, the main bin will contain the values in ascending order.

Create a class `RadixSort` that takes in a List of Integers during object instantiation. The class should have the following properties:
- `items`: is a List of Integers containing the numbers.

It should also have the following methods:
- `sort()`: which returns the sorted numbers from `items` as an `list` of Integers.
- `max_digit()`: which returns the maximum number of digits of all the numbers in `items`. For example, if the numbers are 101, 3, 1041, this method returns 4 as the result since the maximum digit is four from 1041. 
- `convert_to_str(items)`: which returns items as a list of Strings (instead of Integers). This function should pad the higher digits with 0 when converting an Integer to a String. For example if the maximum digit is 4, the following items are converted as follows. From `[101, 3, 1041]` to `["0101", "0003", "1041"]`.

Hint: Your implementation should make use of the generic `Queue` class, which you created, for the bins.

```python
import math
class RadixSort:
    
    def __init__(self, MyList):
        self.items = MyList
        # generate 10 bins and 1 main bin, dictionary comprehension one-liner
        # bin 10 is the main bin
        self.bins = {str(i): Queue() for i in range(11)}
    
    def max_digit(self):
        '''
        to return max digits of the biggest number in list 
        easier soln: 
        return(len(str(max(self.items))))
        '''
        max = -math.inf
        for number in self.items:
            digits = 0
            # if negative, make positive
            if number < 0:
                number = number * -1
            # while number is still divisible, digits += 1
            while number != 0:
                number = number//10
                digits += 1
            # if number of digits is greater than max digits, replace max
            if digits > max:
                max = digits
        
        return max

    def convert_to_str(self, items):
        max_digit = self.max_digit()
        # list comprehension
        return [self.convert(str(number), max_digit) for number in self.items]
            
    def convert(self, number, max_digit):
        return "0" * (max_digit - len(number)) + number

    
    def sort(self):
        # to make my life easier
        main_bin = self.bins['10']

        digit = self.max_digit() - 1
        '''
        pseudocode says to put everything into main bin, but this is quite redundant, so this is deprecated:
        for number in self.convert_to_str(self.items):
            self.bins['10'].enqueue(number)
        '''
        # instead, put number directly into radix bins for first cycle
        for number in self.convert_to_str(self.items):
            self.bins[number[digit]].enqueue(number)
        # empty the radix bins into main bin
        self.empty_bins()
        digit -= 1

        # cycle through digits (index values) from max_digit to 0
        while digit >= 0:
            # if main_bin is emptied out, empty radix bins into main bin and go next digit
            if main_bin.is_empty:
                self.empty_bins()
                digit -= 1
                # to break early once all digits have been cycled through
                if digit == -1:
                    break

            # dequeue from main_bin, and enqueue into radix bin 
            number = main_bin.dequeue()
            self.bins[number[digit]].enqueue(number)
        
        # create new sorted_list from main_bin
        sorted_list = []
        while not main_bin.is_empty:
            sorted_list.append(int(main_bin.dequeue().lstrip('0')))

        self.items = sorted_list
        return self.items     

    def empty_bins(self):
        for i in range(10):
            # going from bin 0 to 9, enqueue everything into main bin
            while not self.bins[str(i)].is_empty:
                self.bins['10'].enqueue(self.bins[str(i)].dequeue())

```

<!-- grade_id: hw51, points: 1, grade: true, solution: false -->
```python
list1 = RadixSort([101, 3, 1041])
assert list1.items == [101,3,1041]
assert list1.max_digit() == 4
assert list1.convert_to_str(list1.items) == ["0101", "0003", "1041"]
ans = list1.sort()
assert ans == [3, 101, 1041]
list2 = RadixSort([23, 1038, 8, 423, 10, 39, 3901])

assert list2.sort() == [8, 10, 23, 39, 423, 1038, 3901]
```

<!-- grade_id: hw52, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python

```