```python
### Grading script code 
### You don't need to read this, proceed to the next cell
import sys
import functools
ipython = get_ipython()

def set_traceback(val):
    method_name = "showtraceback"
    setattr(
        ipython,
        method_name,
        functools.partial(
            getattr(ipython, method_name),
            exception_only=(not val)
        )
    )

class AnswerError(Exception):
  def __init__(self, message):
    pass

def exec_test(f, question):
    try:
        f()
        print(question + " Pass")
    except:
        set_traceback(False) # do not remove
        raise AnswerError(question + " Fail")
```

# Week 2 Problem Set

## Cohort Sessions

**CS1.** *Binary Heap:* Write the following functions:



- `parent_of(index)`: returns the index of node's parent
- `left_of(index)`: returns the index of node's left child
- `right_of(index)`: returns the index of node's right child
- `max_child(array, index, heap_size)`: returns the index of node's largest child. You can assume that the node has at least one child.

Hint: 
- `index` starts from 0. 
- You can refer to the pseudocode in [Binary Heap and Heapsort](https://github.com/Data-Driven-World/d2w_notes/blob/master/BinaryHeap_Heapsort.ipynb).
- When finding the index of the largest child, consider the following cases:
    1. when the node only has one child
    2. when the node has two children


```python
def parent_of(index):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert parent_of(1) == 0 
assert parent_of(2) == 0
assert parent_of(5) == 2
assert parent_of(6) == 2
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def left_of(index):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
assert left_of(0) == 1
assert left_of(1) == 3
assert left_of(3) == 7
assert left_of(6) == 13
```

<!-- grade_id: cs14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def right_of(index):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs15, points: 1, grade: true, solution: false -->
```python
assert right_of(0) == 2
assert right_of(1) == 4
assert right_of(3) == 8
assert right_of(5) == 12
```

<!-- grade_id: cs16, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def max_child(array, index, heap_size):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs17, points: 1, grade: true, solution: false -->
```python
maxheap = [16, 14, 10, 8, 7, 9, 3, 2, 4, 1] 
assert max_child(maxheap, 0, len(maxheap)) == 1
assert max_child(maxheap, 2, len(maxheap)) == 5
assert max_child(maxheap, 3, len(maxheap)) == 8
assert max_child(maxheap, 1, len(maxheap)) == 3
```

<!-- grade_id: cs18, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Binary Heap:* Write two functions 

- `max_heapify(array, index, size)`: that moves the node down so as to satisfy the heap property.  The first argument is the array that contains the heap. The second argument is an integer index where to start the process of heapifying. The third argument is the size of the heap in the array. This argument will be useful in heapsort algorithm where we take out the elements in the array from the heap. 

Hint: You should make use of `size` argument to determine the last element of the heap in the array rather than `len(array)`.

- `build_max_heap(heap)`: that builds the max heap from any array. This function should make use of `max_heapify()` in its definition.

Hint: You can refer to the pseudocode in [Binary Heap and Heapsort](https://github.com/Data-Driven-World/d2w_notes/blob/master/BinaryHeap_Heapsort.ipynb) for the above functions.

```python
def max_heapify(array, index, size):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def build_max_heap(array):
    ###
    ### YOUR CODE HERE
    ###

```

<!-- grade_id: cs23, points: 1, grade: true, solution: false -->
```python
array = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]

build_max_heap(array)
print(array)
assert array == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

<!-- grade_id: cs24, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Heapsort:* Implement heapsort algorithm following the pseudocode in [Binary Heap and Heapsort](https://github.com/Data-Driven-World/d2w_notes/blob/master/BinaryHeap_Heapsort.ipynb).



```python
def heapsort(array):
    ###
    ### YOUR CODE HERE
    ###
    pass
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
array = [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
heapsort(array)
print(array)
assert array == [1, 2, 3, 4, 7, 8, 9, 10, 14, 16]
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Measure computational time of Python's built-in sort function by filling the template below. Hint: 
- You will need the function `gen_random_int()` from Week 01 Problem Set.
- Use `sorted(list)` function of Python's list See [Python's Sorting HOW TO Documentation](https://docs.python.org/3.3/howto/sorting.html). 

```python
import time
import random

def run_function(f, x):
    start = time.time()
    f(x)
    end = time.time()
    return end-start

def gen_random_int(number, seed):
    ###
    ### YOUR CODE HERE
    ###

 
time_builtin = []
# set the maximum power for 10^power number of inputs
maxpower = 7

for n in range(1, maxpower + 1):
    # create array for 10^1, 10^2, etc 
    # use seed 100
    array = None
    
    # call run_function with sorted() and array as arguments
    # result = run_function(None, None) 
    
    ###
    ### YOUR CODE HERE
    ###
    
    time_builtin.append(result)


print(time_builtin)
```